# Aula 1- Exercícios - Sintaxe de Python e NumPy Array

```python
# #@title Init Cell - Execute-me!!!
!pip install ipynb_autograde==0.0.30
from ipynb_autograde.main import init_log, validate
init_log()
email_result=!gcloud config get-value account
if not email_result or 'unset' in email_result[0]:
  !gcloud auth login

```

### Ao resolver um exercício, apague o seguinte trecho do código:
```# YOUR CODE HERE
raise NotImplementedError()
```

### 1.1 Escreva uma função em Python que aceite um inteiro (n) de 1 a 9 e compute o valor de: n+nn+nnn
#### Valide os valores de entrada. Para isso, use a instrução raise ValueError
##### Exemplo: n=5. Então n + nn + nnn = 5 + 55 + 555 = 615

```python
# Faca aqui seus testes do Exercício 1.1

```

<!-- grade_id: cell-a2f686fdd79e4f51, points: 0, grade: false, solution: true -->
```python
# Faca aqui sua função do Exercício 1.1
def calculo_n(n):
    """retorne o valor de: n+nn+nnn
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-5412543179b17057, points: 1, grade: true, solution: false -->
```python
# Faca aqui a validacao do Exercicio 1.1
entradas = [[1], [2], [3], [9]]
saidas = [123, 246, 369, 1107]
validate(calculo_n, entradas, lambda x: x, saidas, "1.1")
```

### 1.2 Escreva uma função para retornar o maior item de uma lista de inteiros

```python
# Faca aqui seus testes do Exercício 1.2

```

<!-- grade_id: cell-2e044a5f0541da37, points: 0, grade: false, solution: true -->
```python
# Faca aqui sua função do Exercício 1.2
def maior(lista):
    """retorne o maior item da lista
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-5bd4079b37b6a7f3, points: 1, grade: true, solution: false -->
```python
# Faca aqui a validacao do Exercicio 1.2
entradas = [[1,0,-1,100,7], [-100,-1580,-1,-9100,-77]]
saidas = [100, -1]
validate(maior, entradas, lambda x: x, saidas, "1.2")
```

### 1.3 Escreva uma função para contar a quantidade de strings, presentes numa lista, cujo tamanho é maior que 2 e o primeiro e último caracteres são os mesmos

#### Exemplo:  ['abc', 'xyz', 'aba', '1221']
#### Resultado esperado : 2

#### Dica 1: A função len retorna o tamanho da string. Exemplo: len('abc'), retorna 3
#### Dica 2: Uma string é uma lista de caracteres

```python
# Faca aqui seus testes do Exercício 1.3

```

<!-- grade_id: cell-59d12cc6e067b728, points: 0, grade: false, solution: true -->
```python
# Faca aqui sua função do Exercício 1.3
def contagem(lista):
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
# Faca aqui a validacao do Exercicio 1.3
v1 = contagem(['abc', 'xyz', 'aba', '1221']) == 2
v2 = contagem(['ac', '1111', 'aasdfsdfdda', '12ewwww21']) == 3
v3 = contagem(['', 'xyzaaaaaaaaa', 'aa', 'babababa']) == 0
current_question_output = v1 and v2 and v3
assert current_question_output
validate(current_question_output, "2.3", False)
```

### 1.4 Crie um código python para imprimir o seguinte padrão. 
#### 1 
#### 1 2 
#### 1 2 3 
#### 1 2 3 4 
#### 1 2 3 4 5

<!-- grade_id: cell-61d03fd9d1b0eb4e, points: 1, grade: true, solution: true -->
```python
# Faca aqui seus testes do Exercício 1.4

## Essa questão não tem assert.
## Faça aqui
# YOUR CODE HERE
raise NotImplementedError()
```

### 1.5 Escreva uma função Python para percorrer uma lista e retornar os números multiplos de N. Pare a iteração no ponto em que encontrar um número maior do que K. Não retorne o 1o número maior do que K.
#### Exemplo:  lista = [120, 115, 12, 32, 56, 200, 715, 152, 130, 155, 183, 300], N=5, K=200
#### Resultado Esperado: [120, 115, 200]

```python
# Faca aqui seus testes do Exercício 1.5
lista = [120, 115, 12, 32, 56, 200, 715, 152, 130, 155, 183, 300]
N = 5
K = 200
```

<!-- grade_id: cell-c9aafbecc3672877, points: 0, grade: false, solution: true -->
```python
# Faca aqui sua função do Exercício 1.5
def filtra_lista(lista, N, K):
    """retorne uma lista com todos os números multiplos de N menores do que K, 
       até o primeiro número maior do que K, não inclua este número na lista de resultado.
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
# Faca aqui a validacao do Exercicio 1.5
v1 = filtra_lista([120, 115, 12, 32, 56, 200, 715, 152, 130, 155, 183, 300], 5, 200) == [120, 115, 200]
current_question_output = v1
assert current_question_output
validate(current_question_output, "2.5", False)
```

### 1.6 Dado um número inteiro qualquer, escreva uma função python para retornar a quantidade de caracteres (dígitos) contida neste número.
##### Exemplo:  qtd_digitos(2345653) deve retornar 7

```python
# Faca aqui seus testes do Exercício 1.6

```

<!-- grade_id: cell-e69a127f71184645, points: 0, grade: false, solution: true -->
```python
# Faca aqui sua função do Exercício 1.6
def qtd_digitos(N):
    """retorne a quantidade de dígitos do número N.
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-418c099df61d368b, points: 1, grade: true, solution: false -->
```python
# Faca aqui a validacao do Exercicio 1.6
v1 = qtd_digitos(2345653) == 7
current_question_output = v1
assert current_question_output
validate(current_question_output, "2.6", True)
```

### 1.7 Escreva uma função Python para somar os valores dos numeros presentes numa string de números separados por vírgula.
#### Exemplo:  "1.0,  4.0,  10.0,  8.0,  4.0,  20.0,  23.0"
#### Resultado esperado: 70

```python
# Faca aqui seus testes do Exercício 1.7

```

<!-- grade_id: cell-f6c50b3e165db993, points: 0, grade: false, solution: true -->
```python
# Faca aqui sua função do Exercício 1.7
def soma_string(texto):
    """retorne a soma dos numeros presentes na variavel texto
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-401910203dadc4db, points: 1, grade: true, solution: false -->
```python
# Faca aqui a validacao do Exercicio 1.7
var = "1.0,  4.0,  10.0,  8.0,  4.0,  20.0,  23.0"
v1 = soma_string(var) == 70.0
current_question_output = v1
assert current_question_output
validate(current_question_output, "2.7", True)
```

### 1.8 Crie um código python para imprimir o seguinte padrão. 
#### Dica: Use uma sintaxe de geração de números em ordem decrescente. Por exemplo ```range(5, 0, -1)``` gera um iterator para a sequência 5, 4, 3, 2, 1
#### 5 4 3 2 1 
#### 4 3 2 1 
#### 3 2 1 
#### 2 1 
#### 1

<!-- grade_id: cell-fc8267fa2d7e4c7c, points: 0, grade: false, solution: true -->
```python
# Faca aqui seus testes do Exercício 1.8

### Este exercicio não tem assert
## Faça aqui

# YOUR CODE HERE
raise NotImplementedError()
```

### 1.9 Usando um loop do tipo for, crie uma função python para inverter uma lista.
#### Exemplo:  inverter([120, 115, 12, 32, 56]) deve retornar [56, 32, 12, 115, 120]

```python
# Faca aqui seus testes do Exercício 1.9


```

<!-- grade_id: cell-513f2b547866e072, points: 0, grade: false, solution: true -->
```python
# Faca aqui sua função do Exercício 1.9
def inverter(lista):
    """retorne uma versão invertida da lista passada como argumento
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
# Faca aqui a validacao do Exercicio 1.9
v1 =  inverter([120, 115, 12, 32, 56]) == [56, 32, 12, 115, 120]
current_question_output = v1
v1 = current_question_output
validate(current_question_output, "2.9", True)
```

### 1.10 Escreva uma função python para substituir vírgula por ponto numa string e depois converta os números para float. Retorne a lista de números.
#### Dica: use a função replace: ```texto.replace("," , ".")``` Este código substitui as virgulas por pontos e retorna uma string com o resultado.
#### Exemplo:  "0,20 0,40 0,56 0,99"
#### Resultado esperado: [0.2, 0.4, 0.56, 0.99]

```python
# Faca aqui seus testes do Exercício 1.10


```

<!-- grade_id: cell-e0a4e5deea9187c4, points: 0, grade: false, solution: true -->
```python
# Faca aqui sua função do Exercício 1.10
def convert_to_float(texto):
    """retorne uma lista de numeros do tipo float
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-30d325468b6f6a67, points: 1, grade: true, solution: false -->
```python
# Faca aqui a validacao do Exercicio 1.10
text = "0,20 0,40 0,56 0,99"
v1 = convert_to_float(text) == [0.2, 0.4, 0.56, 0.99]
current_question_output = v1
assert current_question_output
validate(current_question_output, "2.10", True)
```

### 1.11 Escreva uma função em python para imprimir o padrão a seguir. A função deve aceitar como argumento o tamanho da maior string de asteriscos.
#### Dica: Qual o resultado de se multiplicar uma string por um número em python?

``` 
 * 
 * * 
 * * * 
 * * * * 
 * * * * * 
 * * * * 
 * * * 
 * * 
 *
```

<!-- grade_id: cell-cd8f3abd73a5fcc4, points: 0, grade: false, solution: true -->
```python
# Faca aqui seus testes do Exercício 1.11

### Este exercício não tem assert
## Faça aqui

# YOUR CODE HERE
raise NotImplementedError()
```

### 1.12 Escreva uma função que receba como entrada uma string conforme apresentada a seguir e retorne uma lista contendo a população dos 3 Estados brasileiros mais populosos.

```"""Acre=881935; Alagoas=3337357; Amapá=845731; Amazonas=4144597; Ceará=9132078;
 Distrito Federal=3015268; Espírito Santo=4018650; Goiás=7018354; Maranhão=7075181;
 Mato Grosso=3484466; Mato Grosso do Sul=2778986; Minas Gerais=21168791;
 Paraíba=4018127; Paraná=11433957; Pernambuco=9557071; Piauí=3273227; Rio de Janeiro=17264943;
 Rio Grande do Norte=3506853; Rio Grande do Sul=11377239; Rondônia=1777225;
 Roraima=605761; Santa Catarina=7164788; São Paulo=45919049; Sergipe=2298696;
 Tocantins=1572866; Bahia=14873064; Pará=8602865"""```

#### Dica 1: Nesta string, a informação de cada Estado está separada por um ;
#### Dica 3: A função split quebra uma string numa lista baseado no caracter de separação especificado. 
####     Exemplo: ```'1,2,3'.split(',')``` retorna a lista ```['1','2','3']```
#### Dica 4: a função built-in ```.sort``` de uma lista ordena a referida lista ordenada 

```python
# Faca aqui seus testes do Exercício 1.12

lista_estados="""Acre=881935; Alagoas=3337357; Amapá=845731; Amazonas=4144597; Ceará=9132078;\
 Distrito Federal=3015268; Espírito Santo=4018650; Goiás=7018354; Maranhão=7075181;\
 Mato Grosso=3484466; Mato Grosso do Sul=2778986; Minas Gerais=21168791;\
 Paraíba=4018127; Paraná=11433957; Pernambuco=9557071; Piauí=3273227; Rio de Janeiro=17264943;\
 Rio Grande do Norte=3506853; Rio Grande do Sul=11377239; Rondônia=1777225;\
 Roraima=605761; Santa Catarina=7164788; São Paulo=45919049; Sergipe=2298696;\
 Tocantins=1572866; Bahia=14873064; Pará=8602865"""

```

<!-- grade_id: cell-9e2eb40f94d9f5e4, points: 0, grade: false, solution: true -->
```python
# Faca aqui sua função do Exercício 1.12
def top3_estados(lista_texto):
    """retorne os 3 estados de maior população. O parâmetro de entrada deve ser uma string
         conforme apresentado na descrição do exercício.
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
# Faca aqui a validacao do Exercicio 1.12
lista_estados="""Acre=881935; Alagoas=3337357; Amapá=845731; Amazonas=4144597; Ceará=9132078;\
 Distrito Federal=3015268; Espírito Santo=4018650; Goiás=7018354; Maranhão=7075181;\
 Mato Grosso=3484466; Mato Grosso do Sul=2778986; Minas Gerais=21168791;\
 Paraíba=4018127; Paraná=11433957; Pernambuco=9557071; Piauí=3273227; Rio de Janeiro=17264943;\
 Rio Grande do Norte=3506853; Rio Grande do Sul=11377239; Rondônia=1777225;\
 Roraima=605761; Santa Catarina=7164788; São Paulo=45919049; Sergipe=2298696;\
 Tocantins=1572866; Bahia=14873064; Pará=8602865"""

v1 = set(top3_estados(lista_estados)) == {45919049, 21168791, 17264943}
current_question_output = v1
assert current_question_output
validate(current_question_output, "2.12", True)
```

### 1.13 Escreva uma função Python que aceite como entrada uma string de palavras separadas por vírgula e retorne uma string com as palavras em ordem alfabética e separadas por vírgula. 
#### Exemplo:  Para t="hoje,nota,teste,bolsa"
#### Resultado Esperado: "bolsa,hoje,nota,teste"
#### Dica: use as funções built-in ```.split``` e ```.join``` presente nas strings


```python
# Faca aqui seus testes do Exercício 1.13


```

<!-- grade_id: cell-3201ecbb117ca859, points: 0, grade: false, solution: true -->
```python
# Faca aqui sua função do Exercício 1.13
def sort_words(words):
    """retorne uma string ordenada de palavras separadas por vírgula
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
# Faca aqui a validacao do Exercicio 1.13
v1 = sort_words("hoje,nota,teste,bolsa") == "bolsa,hoje,nota,teste"
current_question_output = v1
assert current_question_output
validate(current_question_output, "2.13", True)
```